# [22 Punkte] Primzahlen

Bei Primzahlen handelt es sich um natürliche Zahlen größer 1, die nur durch 1 und sich selbst ohne Rest teilbar sind. In dieser Aufgabe werden wir zwei Ansätze zur Erzeugung und Überprüfung von Primzahlen vergleichen.

## [12 Punkte] Lookup Table: Sieb des Eratosthenes

### Hintergrund
Wenn *alle* Primzahlen bis zu einer bestimmten Zahl $n$ erzeugt werden sollen, bietet sich das sogenannte Sieb des Eratosthenes an. Dieses geht zunächst davon aus, dass alle Zahlen bis auf 0 und 1 Primzahlen sind, und filtert dann in aufsteigender Folge alle Vielfachen der bis dahin kleinsten noch nicht betrachteten Primzahl aus. Dieses Verfahren muss bis einschließlich $\left\lfloor\sqrt n\,\right\rfloor$ wiederholt werden -- danach werden bis $n$ keine neuen Zahlen mehr ausgesiebt.

Das Ergebnis ist ein Array mit $n+1$ Booleans, das an Index $i$ genau dann `True` ist, wenn die Zahl `i` eine Primzahl ist.

### Aufgabe
In Python kann der Algorithmus wie folgt implementiert werden:

1. Erzeuge eine Liste mit `n+1` Elementen, die alle `True` sind.
2. Setze die Elemente 0 und 1 auf `False` (da 0 und 1 keine Primzahlen sind).
3. Iteriere über den Rest des Arrays bis zum Index `int(n**(1/2))`. Wenn der entsprechende Eintrag `True` ist:
    - erzeuge einen `slice` mit allen Vielfachen dieses Eintrages, um die Teilliste des Arrays mit allen Vielfachen der Primzahl zu erhalten;
    - ersetze diese Teilliste mit einer genauso langen Liste mit Einträgen `False`;

Schreiben Sie eine Funktion `prime_sieve` mit einem Parameter `n`, die eine `list` der Länge `n+1` erzeugt, die an Index `i` genau dann `True` ist, wenn `i` eine Primzahl ist, sonst `False`.

### Spezifikationen
1. **[6 Punkte]** Für $n \le 2$ gibt die Funktion `[False, False, True][:n+1]` zurück
2. **[6 Punkte]** Die Funktion "siebt" für $n > 2$ Kompositzahlen in korrekter Weise aus.

<!-- grade_id: cell-20b0710f46f97a6a, points: 0, grade: false, solution: true -->
```python
# This is a code gap. Students can fill it.
# Sieve function
def prime_sieve(n):
    #Edge case handeling if n < 0
    if n < 0:
        return []
        
    # Step 1 and 2 since it is shorter together
    is_prime = [i > 1 for i in range(n+1)]

    # Loop through every n <= 2 <= n**(1/2) 
    for divider in range(2, int(n**(1/2)) + 1):
        # Loop through every value greater than divider looked at
        for num in range(divider + 1, n + 1):
            # If num is divided by divider, then num is not prime
            if num % divider == 0:
                is_prime[num] = False

    # Return array
    return is_prime

```

```python
import pathlib

import pytest
from pytest_nbgrader import loader

loader.Submission.submit(prime_sieve)

args = ['-qq', '-x', '-W', 'ignore::_pytest.warning_types.PytestAssertRewriteWarning', '--cases']
TESTS = pathlib.Path('tests') / 'PrimeSieve'
```

Output:
```
The following submission will be tested:

def prime_sieve(n):
    #Edge case handeling if n < 0
    if n < 0:
        return []
        
    # Step 1 and 2 since it is shorter together
    is_prime = [i > 1 for i in range(n+1)]

    # Loop through every n <= 2 <= n**(1/2) 
    for divider in range(2, int(n**(1/2)) + 1):
        # Loop through every value greater than divider looked at
        for num in range(divider + 1, n + 1):
            # If num is divided by divider, then num is not prime
            if num % divider == 0:
                is_prime[num] = False

    # Return array
    return is_prime
```

<!-- grade_id: cell-be4cd2ef04795568, points: 6, grade: true, solution: false -->
```python
assert pytest.main([*args, TESTS / 'correctness_small_n.yml']) is pytest.ExitCode.OK
```

Output:
```
[33ms[0m[32m.[0m[32m.[0m[32m.[0m[32m                                                                     [100%][0m
```

<!-- grade_id: cell-372f5bfd90f59bc7, points: 6, grade: true, solution: false -->
```python
assert pytest.main([*args, TESTS / 'correctness_medium_n.yml']) is pytest.ExitCode.OK
```

Output:
```
[33ms[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m [ 73%]
[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[3
```

## [10 Punkte] Funktion: Trial Division

### Hintergrund
Wenn nur *einzelne* Zahlen darauf überprüft werden sollen, ob sie Primzahlen sind, lohnt es sich nicht, mit dem Sieb des Eratosthenes *alle* Primzahlen bis zu dieser Zahl zu erzeugen. Stattdessen reicht es, nur für die eingegebene Zahl $n$ nacheinander zu überprüfen, ob diese Zahl durch irgendeine der Zahlen in der Menge $\{2, 3, ..., \left\lfloor \sqrt{n} \right\rfloor\}$ teilbar ist.

### Aufgabe
Schreiben Sie eine Funktion `is_prime` mit genau einem Parameter `n`. Die Funktion soll `True` zurückgeben, falls diese Zahl eine Primzahl ist, sonst `False`.

### Spezifikationen
1. **[5 Punkte]** Die Funktion gibt korrekt zurück, dass 0 und 1 keine Primzahlen sind, 2 hingegen schon.
2. **[5 Punkte]** Für größere $n$ gibt die Funktion korrekt aus, ob $n$ eine Primzahl ist.

<!-- grade_id: cell-c26ecb9b8e848643, points: 0, grade: false, solution: true -->
```python
# This is a code gap. Students can fill it.
# Is prime function
def is_prime(n):
    # Edge case: n is 1 or lower
    if n < 2:
        return False

    # Check all numbers in [2,sqr(n)]
    for divisor in range(2, int(n**(1/2)) + 1):
        # If n is divided by any number in range it's not prime
        if n % divisor == 0:
            return False

    # If all checks were not divisors then it is prime
    return True
```

```python
import pathlib

import pytest
from pytest_nbgrader import loader

loader.Submission.submit(is_prime)

args = ['-qq', '-x', '-W', 'ignore::_pytest.warning_types.PytestAssertRewriteWarning', '--cases']
TESTS = pathlib.Path('tests') / 'TrialDivision'
```

Output:
```
The following submission will be tested:

def is_prime(n):
    # Edge case: n is 1 or lower
    if n < 2:
        return False

    # Check all numbers in [2,sqr(n)]
    for divisor in range(2, int(n**(1/2)) + 1):
        # If n is divided by any number in range it's not prime
        if n % divisor == 0:
            return False

    # If all checks were not divisors then it is prime
    return True
```

<!-- grade_id: cell-c551219334432ca1, points: 5, grade: true, solution: false -->
```python
assert pytest.main([*args, TESTS / 'correctness_small_n.yml', '-q', '-x']) is pytest.ExitCode.OK
```

Output:
```
[33ms[0m[32m.[0m[32m.[0m[32m.[0m[32m                                                                     [100%][0m
```

<!-- grade_id: cell-d9137804f58da66f, points: 5, grade: true, solution: false -->
```python
assert pytest.main([*args, TESTS / 'correctness_medium_n.yml', '-q', '-x']) is pytest.ExitCode.OK
```

Output:
```
[33ms[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m [ 73%]
[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[3
```